---
title: Test
description: >-
  Test


---
## mpc

```yaml
type: MultipleChoiceExercise

xp: NaN

key: 88d71c436a
```

here is the assignment / question

`@instructions`
- 1
- 2
- 3

`@hint`
this is the hint

`@pre_exercise_code`
```{undefined}
from math import pi
```


`@sct`
```{undefined}
# this is where the code goes for sct
```




