---
title: Test
description: >-
  Test


---
## Sample exercise

```yaml
type: NormalExercise
lang: python
xp: 100
skills: 2
key: d50a58374f
```














---
## mpc

```yaml
type: MultipleChoiceExercise

xp: 

key: 88d71c436a
```

here is the assignment / question

`@instructions`
- 1
- 2
- 3

`@hint`
this is the hint

`@pre_exercise_code`
```{}
from math import pi
```


`@sct`
```{}
# this is where the code goes for sct
```




