---
title: Test
description: >-
  Test


---
## Sample exercise

```yaml
type: NormalExercise
lang: python
xp: 100
skills: 2
key: d50a58374f
```














---
## Insert exercise title here

```yaml
type: NormalExercise

xp: 

key: 5bbdffad13
```

this is the assignment for normal exercise title

`@instructions`
those are the instructions

`@hint`
this is the hint

`@pre_exercise_code`
```{}
from math import pi
```
`@sample_code`
```{}
# Starting point
```
`@solution`
```{}
print(pi)
```
`@sct`
```{}
# here we write sct's
```





---
## mpc

```yaml
type: MultipleChoiceExercise

xp: 

key: 88d71c436a
```

here is the assignment / question

`@instructions`
- 1
- 2
- 3

`@hint`
this is the hint

`@pre_exercise_code`
```{}
from math import pi
```


`@sct`
```{}
# this is where the code goes for sct
```




